---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: c909bf2ad2
lang: r
xp: 100
skills: 1
```

The purpose of this exercise is to illustrate the basic elements of forecasting. We will use google sheets INTERCEP, SLOPE and RSQ functions initially and then move on to using Trendlines

`@instructions`
- Watch the video with the demonstrations
- Complete the Trendline Exercise mentioned on the Line by Line Predictions/Forecasts slide
- Practice with the amazon.com Sales, Total Assets and Net Income Data

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}
# Load datasets and packages here.
```

`@sample_code`
```{r}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{r}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Forecasting Sales Using Trendline and Similar Techniques

```yaml
type: VideoExercise
key: 68b4293e5b
xp: 50
```

`@projector_key`
33e68927c43d9191475cd2b038073a8c
